<problem display_name="MatrixGrader" markdown="null">
<script type="text/python">
from mitxgraders import *

grader = MatrixGrader(
  answers='[[2*x, 4], [-2, 6*x]]',
  variables=['x'],
  max_array_dim=2
)
</script>

<p>Enter the following matrix in the input box:</p>
  


<p>Suggested inputs:</p>

<ul>
  <li><code>[[2*x, 4], [-2, 6*x]]</code></li>
  <li><code>2*[[x, 2], [-1, 3*x]]</code></li>
  <li><code>2*[[0, 2], [-1, 0]]+2*x*[[1, 0], [0, 3]]</code></li>
  <li>Try inputting ill-formed matrices such as <code>[[1, 2], [3]]</code></li>
</ul>

<customresponse cfn="grader" expect="[[2*x, 4], [-2, 6*x]]">
  <textline math="true" size="40"/>
</customresponse>
</problem>
